.. activecode:: ac13_6_9
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: PythonModules
   :subchapter: ChapterAssessment
   :topics: PythonModules/ChapterAssessment
   :from_source: F

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(test, True, "Testing that test has the correct value.")

   myTests().main()