.. activecode:: ac13_6_8
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: PythonModules
   :subchapter: ChapterAssessment
   :topics: PythonModules/ChapterAssessment
   :from_source: F


   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(top, 3, "Testing that top has the correct value.")
      def testTwo(self):
         self.assertEqual(bottom, 94, "Testing that bottom has the correct value.")


   myTests().main()