.. activecode:: ac13_6_6
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: PythonModules
   :subchapter: ChapterAssessment
   :topics: PythonModules/ChapterAssessment
   :from_source: F

   a = 5
   b = 8
   d = 125

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testA(self):
         self.assertEqual(c, 390625, "Testing that c has the correct value.")
      def testB(self):
         self.assertEqual(e, 3125, "Testing that e has the correct value.")

   myTests().main()