.. actex:: ac17_5_1
   :author: bmiller
   :difficulty: 1.4250899179
   :basecourse: fopp
   :chapter: NestedData
   :subchapter: Exercises
   :topics: NestedData/Exercises
   :from_source: T
   :pct_on_first: 0.0762711864
   :total_students_attempting: 118
   :num_students_correct: 83.0
   :mean_clicks_to_correct: 8.2168674699

   Iterate through the list so that if the character 'm' is in the string, then it should be added to a new list called ``m_list``. Hint: Because this isn't just a list of lists, think about what type of object you want your data to be stored in. Conditionals may help you.
   ~~~~
   
   d = ['good morning', 'hello', 'chair', 'python', ['music', 'flowers', 'facebook', 'instagram', 'snapchat', ['On my Own', 'monster', 'Words dont come so easily', 'lead me right']], 'Stressed Out', 'Pauver Coeur', 'Reach for Tomorrow', 'mariners song', 'Wonder sleeps here']
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(m_list, ['good morning', 'music', 'instagram', 'On my Own', 'monster', 'Words dont come so easily', 'lead me right', 'Reach for Tomorrow', 'mariners song'], "Testing that m_list has the correct list assigned to it.")
   
   myTests().main()