.. activecode:: ac10_3_7
    :author: bmiller
    :difficulty: 1.0404826638
    :basecourse: fopp
    :chapter: Dictionaries
    :subchapter: Dictionarymethods
    :topics: Dictionaries/Dictionarymethods
    :from_source: T
    :language: python
    :autograde: unittest
    :practice: T
    :pct_on_first: 0.5101180438
    :total_students_attempting: 1186
    :num_students_correct: 873.0
    :mean_clicks_to_correct: 1.6872852234

    **5.** We have a dictionary of the specific events that Italy has won medals in and the number of medals they have won for each event. Assign to the variable ``events`` a list of the keys from the dictionary ``medal_events``. Do not hard code this.
    
    ~~~~
    
    
    medal_events = {'Shooting': 7, 'Fencing': 4, 'Judo': 2, 'Swimming': 3, 'Diving': 2}
    
    =====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def testOne(self):
            self.assertEqual(type(events), list, "Testing that events is a list")
            self.assertEqual(sorted(events), sorted(medal_events), "Testing that events was created correctly")
            self.assertNotIn('[', self.getEditorText(), "No hard coding")
    
    myTests().main()