.. activecode:: ac18_7_2
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: Sorting
   :subchapter: ChapterAssessment
   :topics: Sorting/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T

   Sort the list below, ``animals``, into alphabetical order, a-z. Save the new list as ``animals_sorted``.
   ~~~~

   animals = ['elephant', 'cat', 'moose', 'antelope', 'elk', 'rabbit', 'zebra', 'yak', 'salamander', 'deer', 'otter', 'minx', 'giraffe', 'goat', 'cow', 'tiger', 'bear']

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(animals_sorted, ['antelope', 'bear', 'cat', 'cow', 'deer', 'elephant', 'elk', 'giraffe', 'goat', 'minx', 'moose', 'otter', 'rabbit', 'salamander', 'tiger', 'yak', 'zebra'], "Testing that animals_sorted was created correctly.")

   myTests().main()