.. activecode:: ac5_5_5
   :author: bmiller
   :difficulty: 1.0300088564
   :basecourse: fopp
   :chapter: Sequences
   :subchapter: Length
   :topics: Sequences/Length
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.5879494655
   :total_students_attempting: 2058
   :num_students_correct: 1637.0
   :mean_clicks_to_correct: 1.50946854

   Assign the number of elements in ``lst`` to the variable ``output``.
   ~~~~
   lst = ["hi", "morning", "dog", "506", "caterpillar", "balloons", 106, "yo-yo", "python", "moon", "water", "sleepy", "daffy", 45, "donald", "whiteboard", "glasses", "markers", "couches", "butterfly", "100", "magazine", "door", "picture", "window", ["Olympics", "handle"], "chair", "pages", "readings", "burger", "juggle", "craft", ["store", "poster", "board"], "laptop", "computer", "plates", "hotdog", "salad", "backpack", "zipper", "ring", "watch", "finger", "bags", "boxes", "pods", "peas", "apples", "horse", "guinea pig", "bowl", "EECS"]
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testTwo(self):
         self.assertEqual(output, 52, "Testing that output value is assigned to correct value.")
   
   myTests().main()