.. activecode:: ac10_2_1
   :author: bmiller
   :difficulty: 1.0148891864
   :basecourse: fopp
   :chapter: Dictionaries
   :subchapter: Dictionaryoperations
   :topics: Dictionaries/Dictionaryoperations
   :from_source: T
   :language: python
   :autograde: unittest
   :pct_on_first: 0.7005307051
   :total_students_attempting: 1319
   :num_students_correct: 1080.0
   :mean_clicks_to_correct: 1.2527777778

   **2.** Update the value for "Phelps" in the dictionary ``swimmers`` to include his medals from the Rio Olympics by adding 5 to the current value (Phelps will now have 28 total medals). Do not rewrite the dictionary.
   ~~~~
   
   swimmers = {'Manuel':4, 'Lochte':12, 'Adrian':7, 'Ledecky':5, 'Dirado':4, 'Phelps':23}
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testThree(self):
         self.assertEqual(sorted(swimmers.items()), sorted([('Adrian', 7), ('Dirado', 4), ('Ledecky', 5), ('Lochte', 12), ('Phelps', 28), ('Manuel',4)]), "Testing that swimmers is assigned to correct values.")
   
   myTests().main()