.. activecode:: ac5_3_6
   :author: bmiller
   :difficulty: 1.0320702046
   :basecourse: fopp
   :chapter: Sequences
   :subchapter: IndexOperatorWorkingwiththeCharactersofaString
   :topics: Sequences/IndexOperatorWorkingwiththeCharactersofaString
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.6249395843
   :total_students_attempting: 2069
   :num_students_correct: 1653.0
   :mean_clicks_to_correct: 1.5444646098

   Assign the value of the last chacter of ``lst`` to the variable ``output``. Do this so that the length of lst doesn't matter.
   ~~~~
   lst = "Every chess or checkers game begins from the same position and has a finite number of moves that can be played. While the number of possible scenarios and moves is quite large, it is still possible for computers to calculate that number and even be programmed to respond well against a human player..."
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testThree(self):
         self.assertEqual(output, ".", "Testing that output value is assigned to correct value.")
   
   myTests().main()