.. activecode:: ac9_10_7
   :author: bmiller
   :difficulty: 1.0411052919
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: travel_plans.txt
   :nocodelens: 
   :pct_on_first: 0.5537065053
   :total_students_attempting: 661
   :num_students_correct: 513.0
   :mean_clicks_to_correct: 1.6978557505

   Assign the first 33 characters from the textfile, ``travel_plans.txt`` to the variable ``first_chars``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testFive(self):
         self.assertEqual(first_chars, "This summer I will be travelling.", "Testing that first_chars is assigned to correct value.")
   
   myTests().main()