.. activecode:: ac9_10_1
   :author: bmiller
   :difficulty: 1.0713199539
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: travel_plans.txt
   :nocodelens: 
   :pct_on_first: 0.488
   :total_students_attempting: 875
   :num_students_correct: 721.0
   :mean_clicks_to_correct: 2.2108183079

   The textfile, ``travel_plans.txt``, contains the summer travel plans for someone with some commentary. Find the total number of characters in the file and save to the variable ``num``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(num, 316, "Testing that num value is assigned to correct value.")
   
   myTests().main()