.. activecode:: ac9_10_8
   :author: bmiller
   :difficulty: 1.2328306938
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: school_prompt.txt
   :nocodelens: 
   :pct_on_first: 0.325613079
   :total_students_attempting: 734
   :num_students_correct: 530.0
   :mean_clicks_to_correct: 4.9528301887

   **Challenge:** Using the file ``school_prompt.txt``, if the character 'p' is in a word, then add the word to a list called ``p_words``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(p_words, ['topic', 'point', 'papers,', 'ups,', 'scripts.'], "Testing that p_words has the correct list.")
   
   myTests().main()