.. activecode:: ac9_10_4
   :author: bmiller
   :difficulty: 1.0801109512
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: school_prompt.txt
   :nocodelens: 
   :pct_on_first: 0.4907161804
   :total_students_attempting: 754
   :num_students_correct: 611.0
   :mean_clicks_to_correct: 2.3600654664

   Assign the first 30 characters of ``school_prompt.txt`` as a string to the variable ``beginning_chars``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(len(beginning_chars), 30, "Testing that beginning_chars has the correct length.")
         self.assertEqual(beginning_chars, "Writing essays for school can ", "Testing that beginning_chars has the correct string.")
   
   myTests().main()