.. activecode:: ac9_10_2
   :author: bmiller
   :difficulty: 1.2662729349
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: emotion_words.txt
   :nocodelens: 
   :pct_on_first: 0.2703016241
   :total_students_attempting: 862
   :num_students_correct: 680.0
   :mean_clicks_to_correct: 5.5205882353

   We have provided a file called ``emotion_words.txt`` that contains lines of words that describe emotions. Find the total number of words in the file and assign this value to the variable ``num_words``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(num_words, 48, "Testing that num_words was assigned to the correct value.")
   
   myTests().main()