.. activecode:: ac9_10_6
   :author: bmiller
   :difficulty: 1.0916631606
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: emotion_words.txt
   :nocodelens: 
   :pct_on_first: 0.4724637681
   :total_students_attempting: 690
   :num_students_correct: 525.0
   :mean_clicks_to_correct: 2.5561904762

   **Challenge:** Create a list called ``emotions`` that contains the first word of every line in ``emotion_words.txt``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(emotions, ['Sad', 'Angry', 'Happy', 'Confused', 'Excited', 'Scared', 'Nervous'], "Testing that emotions was created correctly.")
   
   myTests().main()