.. activecode:: ac9_10_3
   :author: bmiller
   :difficulty: 1.0362741288
   :basecourse: fopp
   :chapter: Files
   :subchapter: ChapterAssessment
   :topics: Files/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :available_files: school_prompt.txt
   :nocodelens: 
   :pct_on_first: 0.6147239264
   :total_students_attempting: 815
   :num_students_correct: 682.0
   :mean_clicks_to_correct: 1.6158357771

   Assign to the variable ``num_lines`` the number of lines in the file ``school_prompt.txt``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(num_lines, 10, "Testing that num_lines has the correct value.")
   
   myTests().main()