.. activecode:: question10_7_1
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 10-tuples
   :subchapter: 10-usingtuplesaskeys
   :topics: 10-tuples/10-usingtuplesaskeys
   :from_source: T
   :practice: T
   :nocodelens:

   Write code to create a dictionary called 'd1', and in it give the tuple (1, "a") a value of "tuple".
   ~~~~

   ====
   from unittest.gui import TestCaseGui

   class MyTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(d1[(1, "a")], "tuple", "Checking to make sure the key has the correct value.")

   MyTests().main()