.. activecode:: ac17_2_1
   :author: bmiller
   :difficulty: 1.0906408807
   :basecourse: fopp
   :chapter: NestedData
   :subchapter: NestedDictionaries
   :topics: NestedData/NestedDictionaries
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.3164179104
   :total_students_attempting: 335
   :num_students_correct: 206.0
   :mean_clicks_to_correct: 2.5388349515

   **1.** Extract the value associated with the key color and assign it to the variable ``color``. Do not hard code this.
   
   ~~~~
   
   info = {'personal_data':
            {'name': 'Lauren',
             'age': 20,
             'major': 'Information Science',
             'physical_features':
                {'color': {'eye': 'blue',
                           'hair': 'brown'},
                 'height': "5'8"}
            },
          'other':
            {'favorite_colors': ['purple', 'green', 'blue'],
             'interested_in': ['social media', 'intellectual property', 'copyright', 'music', 'books']
            }
         }
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(color, {'eye': 'blue', 'hair': 'brown'}, "Testing that color has the correct value.")
   
   myTests().main()