.. activecode:: assess_ac_4_1_3_2
    :author: bmiller
    :difficulty: 3.0
    :basecourse: fopp
    :chapter: TransformingSequences
    :subchapter: week4a1
    :topics: TransformingSequences/week4a1
    :from_source: T
    :language: python
    :autograde: unittest
    :practice: T
    :topics: TransformingSequences/MutatingMethods

    Write code to find the position of the string "Tony" in the list ``awards`` and save that information in the variable ``pos``.
    ~~~~
    awards = ['Emmy', 'Tony', 'Academy', 'Grammy']


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_output(self):
        self.assertEqual(pos, 1, "Testing that pos is set correctly.")
        self.assertIn('.index(', self.getEditorText(), "Testing that a method invocation was used in your code.")

    myTests().main()