.. activecode:: transforming_sequences_irma_q3
   :author: Irma Ravkic
   :difficulty: 0.0
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: week4a1
   :topics: TransformingSequences/week4a1
   :from_source: F
   :language: python
   :autograde: unittest
   :practice: T
   :topics: TransformingSequences/TheAccumulatorPatternwithStrings

   Below are a set of scores that students have received in the past semester. Write code to determine the average score and put it in a variable ``avg_score``.
   ~~~~
   scores = "67 80 90 78 93 20 79 89 96 97 92 88 79 68 58 90 98 100 79 74 83 88 80 86 85 70 90 100"

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertAlmostEqual(avg_score, 82.0357142857, 7, "Testing that you got the right average.")
         self.assertIn('for', self.getEditorText(), "Testing that you used a for loop.")

   myTests().main()