.. activecode:: act_kiva_9
    :author: bmiller
    :difficulty: 3.0
    :basecourse: fopp
    :chapter: Projects
    :subchapter: kiva_statistics
    :topics: Projects/kiva_statistics
    :from_source: T
    :include: act_kiva_1
    :pct_on_first: 0.3246753247
    :total_students_attempting: 616
    :num_students_correct: 528.0
    :mean_clicks_to_correct: 2.9090909091

    What is the name of the country with the loan that took the longest to raise?  Store your result in the variable ``longest_to_fund``
    ~~~~
    # Your code here
    ====
    from unittest.gui import TestCaseGui
    
    class MyTests(TestCaseGui):
    
        def testOne(self):
            self.assertTrue('longest_to_fund' in self.getEditorText(), "you need a loan_total variable")
            self.assertEqual(longest_to_fund, country_name[time_to_raise.index(max(time_to_raise))], "Hint: max and index are your friends")
    
    MyTests().main()