.. activecode:: lhs_12_11
    :author: LHS CS Team
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: Dictionaries
    :subchapter: Exercises
    :topics: Dictionaries/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 0.6699507389
    :total_students_attempting: 203
    :num_students_correct: 196.0
    :mean_clicks_to_correct: 2.2602040816

    Provided is the dictionary, ``medal_count``, which lists countries and their
    respective medal count at the halfway point in the 2016 Rio Olympics.
    Using dictionary mechanics, assign the medal count value for ``"Belarus"``
    to the variable ``belarus``. Do not hardcode this.
    
    ~~~~   
    medal_count = {'United States': 70, 'Great Britain':38, 'China':45, 'Russia':30, 'Germany':17, 'Italy':22, 'France': 22, 'Japan':26, 'Australia':22, 'South Korea':14, 'Hungary':12, 'Netherlands':10, 'Spain':5, 'New Zealand':8, 'Canada':13, 'Kazakhstan':8, 'Colombia':4, 'Switzerland':5, 'Belgium':4, 'Thailand':4, 'Croatia':3, 'Iran':3, 'Jamaica':3, 'South Africa':7, 'Sweden':6, 'Denmark':7, 'North Korea':6, 'Kenya':4, 'Brazil':7, 'Belarus':4, 'Cuba':5, 'Poland':4, 'Romania':4, 'Slovenia':3, 'Argentina':2, 'Bahrain':2, 'Slovakia':2, 'Vietnam':2, 'Czech Republic':6, 'Uzbekistan':5}
    
    # your code here to assign the medal count value for Belarus to the variable belarus
    # from the medal_events dictionary.
        
    ====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def testSix(self):
            self.assertEqual(belarus, 4, "Testing that belarus is assigned the correct value.")
    
    myTests().main()