.. activecode:: LHSQuestionThree
    :author: Wesley Thang
    :difficulty: 1.0
    :basecourse: thinkcspy
    :chapter: Dictionaries
    :subchapter: Exercises
    :topics: Dictionaries/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 1.0
    :total_students_attempting: 1
    :num_students_correct: 1.0
    :mean_clicks_to_correct: 1.0

    You are keeping track of olympic medals for Italy in the 2016 Rio Summer Olympics!
    At the moment, Italy has 7 gold medals, 8 silver metals, and 6 bronze medals.
    Create a dictionary called ``olympics`` where the keys are the types of medals,
    and the values are the number of that type of medals that Italy has won so far. 
    
    ~~~~       
    # your code here       
    ====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def testOne(self):
            self.assertEqual(sorted(olympics.items()), sorted([('gold', 7), ('silver', 8), ('bronze', 6)]), "Testing that olympics was created correctly.")     
    
    myTests().main()