.. actex:: ac8_11_1
   :author: bmiller
   :difficulty: 1.1758619176
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: Exercises
   :topics: TransformingSequences/Exercises
   :from_source: T
   :pct_on_first: 0.3328025478
   :total_students_attempting: 628
   :num_students_correct: 488.0
   :mean_clicks_to_correct: 3.9856557377

   For each word in the list ``verbs``, add an -ing ending. Overwrite the old list so that ``verbs`` has the same words with ``ing`` at the end of each one.
   ~~~~
   verbs = ["kayak", "cry", "walk", "eat", "drink", "fly"]
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testSeven(self):
           self.assertEqual(verbs, ['kayaking', 'crying', 'walking', 'eating', 'drinking', 'flying'], "Testing that verbs is assigned to correct values.")
   
   myTests().main()