.. activecode:: transforming_seq_q_1
   :author: Irma Ravkic
   :difficulty: 1.0
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: Exercises
   :topics: TransformingSequences/Exercises
   :from_source: F
   :language: python
   :autograde: unittest
   :pct_on_first: 0.0974025974
   :total_students_attempting: 154
   :num_students_correct: 87.0
   :mean_clicks_to_correct: 1.9770114943

   
   Consider the code given to you below. It iterates through a list of 5 integers, and appends them to a list. Run the code first and see the result. 
   Now, by only modifying line 3 change the code to instead of just adding an element to the list, it adds each number multiplied by two.
   ~~~~
   my_list = []
   for num in range(1, 6):
      my_list.append(num)
   print(my_list)
   ====
   from unittest.gui import TestCaseGui
   import re
   class myTests(TestCaseGui):
       def testOne(self):
           self.assertEqual('[2, 4, 6, 8, 10]\n', self.getOutput(), 'Checking output.')
   
   myTests().main()