.. activecode:: ac14_2_2
   :author: bmiller
   :difficulty: 1.1309780682
   :basecourse: fopp
   :chapter: MoreAboutIteration
   :subchapter: ThewhileStatement
   :topics: MoreAboutIteration/ThewhileStatement
   :from_source: T
   :practice: T
   :pct_on_first: 0.3487084871
   :total_students_attempting: 542
   :num_students_correct: 389.0
   :mean_clicks_to_correct: 3.2236503856

   Write a while loop that is initialized at 0 and stops at 15. If the counter is an even number, append the counter to a list called ``eve_nums``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
     def testOne(self):
        self.assertEqual(eve_nums, [0,2,4,6,8,10,12,14], "Testing that eve_nums has been assigned the correct elements")
   
   myTests().main()