.. activecode:: ac14_2_3
    :author: bmiller
    :difficulty: 1.1195069295
    :basecourse: fopp
    :chapter: MoreAboutIteration
    :subchapter: ThewhileStatement
    :topics: MoreAboutIteration/ThewhileStatement
    :from_source: T
    :practice: T
    :pct_on_first: 0.2408088235
    :total_students_attempting: 544
    :num_students_correct: 346.0
    :mean_clicks_to_correct: 3.0289017341

    Below, we've provided a for loop that sums all the elements of ``list1``. Write code that accomplishes the same task, but instead uses a while loop. Assign the accumulator variable to the name ``accum``.
    ~~~~
    
    list1 = [8, 3, 4, 5, 6, 7, 9]
    
    tot = 0
    for elem in list1:
        tot = tot + elem
    
    =====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
       def testOne(self):
          self.assertEqual(accum, 42, "Testing that accum has the correct value.")
          self.assertIn('while', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
    
    
    myTests().main()