.. activecode:: ac12_3_3
   :author: bmiller
   :difficulty: 1.0495178454
   :basecourse: fopp
   :chapter: Tuples
   :subchapter: TuplesasReturnValues
   :topics: Tuples/TuplesasReturnValues
   :from_source: T
   :language: python
   :autograde: unittest
   :chatcodes: 
   :practice: T
   :pct_on_first: 0.5880829016
   :total_students_attempting: 386
   :num_students_correct: 295.0
   :mean_clicks_to_correct: 1.8406779661

   Define a function called ``information`` that takes as input, the variables ``name``, ``birth_year``, ``fav_color``, and ``hometown``. It should return a tuple of these variables in this order.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(information("Sarah", 1996, "purple", "St. Louis"), ("Sarah", 1996, "purple", "St. Louis"), "Testing that information returns the correct tuple on input ('Sarah', 1996, 'purple', 'St. Louis')")
   
   myTests().main()