.. activecode:: itr-ex-countdownq
    :author: bmiller
    :difficulty: 3.0
    :basecourse: py4e-int
    :chapter: 05-iterations
    :subchapter: 05-WriteCode
    :topics: 05-iterations/05-WriteCode
    :from_source: T
    :practice: T
    :autograde: unittest

    def countdown():
        counter = 10
        while Counter > 10:
            Print(counter)
            counter = counter + 1

    countdown()

    ====
    from unittest.gui import TestCaseGui
    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(recPerimeter(10, 20),60,"Tested recPerimeter on inputs 10 and 20")
            self.assertEqual(recPerimeter(1, 2),6,"Tested recPerimeter on inputs 1 and 2")
            self.assertEqual(recPerimeter(23.2, 12.3),71,"Tested recPerimeter on inputs 23.2 and 12.3")
            self.assertEqual(recPerimeter(3.0, 2),10.0,"Tested recPerimeter on inputs 3.0 and 2")

    myTests().main()