.. activecode:: itr-ex-allValuesq
    :author: bmiller
    :difficulty: 3.0
    :basecourse: py4e-int
    :chapter: 05-iterations
    :subchapter: 05-WriteCode
    :topics: 05-iterations/05-WriteCode
    :from_source: T
    :practice: T
    :autograde: unittest

    Finish lines 1 and 5 so that the following code correct prints all the values from -5 to -1.
    ~~~~
    output =
    x = -5
    while x < 0:
        output = output + str(x) + " "
        x =
    print(output)

    ====
    from unittest.gui import TestCaseGui
    class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(output,"-5 -4 -3 -2 -1 ")

    myTests().main()