.. activecode:: ac6_6_4
   :author: bmiller
   :difficulty: 1.1397673381
   :basecourse: fopp
   :chapter: Iteration
   :subchapter: TheAccumulatorPattern
   :topics: Iteration/TheAccumulatorPattern
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.2520515826
   :total_students_attempting: 1706
   :num_students_correct: 1290.0
   :mean_clicks_to_correct: 3.3728682171

   Write code to create a list of integers from 0 through 52 and assign that list to the variable ``numbers``. You should use a special Python function -- do not type out the whole list yourself. HINT: You can do this in one line of code!
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(numbers, range(53), "Testing that numbers is a list that contains the correct elements.")
   
   myTests().main()