.. activecode:: writingcode_question9_10
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 09-dictionaries
   :subchapter: 09-writecode
   :topics: 09-dictionaries/09-writecode
   :from_source: T
   :practice: T
   :available_files: romeo.txt

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(len(counts), 26, "Making sure all 26 words made it into the dictionary.")
           self.assertEqual(counts['is'], 3, "Checking 'is' was only counted for three times.")
           self.assertEqual(counts['arise'], 1, "Checking to make sure 'arise' appears once in the dictionary.")

   myTests().main()