.. activecode:: writingcode_question9_9
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 09-dictionaries
   :subchapter: 09-writecode
   :topics: 09-dictionaries/09-writecode
   :from_source: T
   :practice: T
   :available_files: words.txt

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(len(word_count), 119, "Making sure all the words were accounted for.")
           self.assertEqual(word_count['and'], 5, "Checking if 'and' is counted for 5 times.")
           self.assertEqual(word_count['what'], 3, "Checking if 'what' appears 3 times.")

   myTests().main()