.. activecode:: writingcode_question9_6
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 09-dictionaries
   :subchapter: 09-writecode
   :topics: 09-dictionaries/09-writecode
   :from_source: T
   :practice: T
   :available_files: mbox-short.txt

   Write this program to record the domain name (instead of the address) where the message was sent from instead of who the mail came from (i.e., the whole email address). At the end of the program, print out the contents of your dictionary.
   ~~~~

   with open("mbox-short.txt2", "r") as filename:
       message_count = {}

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(message_count['iupui.edu'], '5', "Making sure 'iupui.edu' sent 5 emails.")
           self.assertEqual(len(message_count), 6, "Checking that all the emails made it into the dictionary.")

   myTests().main()