.. activecode:: writingcode_question9_5
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 09-dictionaries
   :subchapter: 09-writecode
   :topics: 09-dictionaries/09-writecode
   :from_source: T
   :practice: T
   :available_files: mbox-short.txt

   with open("mbox-short.txt3", "r") as filename:
       message_count = {}

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(message_count['cwen@iupui.edu'], '5', "Making sure 5 emails were sent from this email address.")
           self.assertEqual(len(message_count), 11, "Checking that all the emails made it into the dictionary.")

   myTests().main()