.. actex:: ac11_14_3
   :author: bmiller
   :difficulty: 1.1254431752
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Exercises
   :topics: Functions/Exercises
   :from_source: T
   :pct_on_first: 0.4403183024
   :total_students_attempting: 377
   :num_students_correct: 347.0
   :mean_clicks_to_correct: 3.1296829971

   Write a function that reverses its string argument.
   ~~~~
   def reverse(astring):
       # your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
           self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
           self.assertEqual(reverse(""),"","Tested reverse on input of ''")
   
   
   
   
   myTests().main()