.. actex:: ac11_14_4
   :author: bmiller
   :difficulty: 1.0591253905
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Exercises
   :topics: Functions/Exercises
   :from_source: T
   :nocodelens: 
   :pct_on_first: 0.5342960289
   :total_students_attempting: 277
   :num_students_correct: 264.0
   :mean_clicks_to_correct: 2.0037878788

   Write a function that mirrors its string argument,
   generating a string containing the original string and the string backwards.
   ~~~~
   
   def mirror(mystr):
       # your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(mirror("good"),"gooddoog","Tested mirror on input of 'good'")
           self.assertEqual(mirror("Python"),"PythonnohtyP","Tested mirror on input of 'Python'")
           self.assertEqual(mirror(""),"","Tested mirror on input of ''")
           self.assertEqual(mirror("a"),"aa","Tested mirror on input of 'a'")
   
   
   myTests().main()