.. actex:: ac11_14_5
   :author: bmiller
   :difficulty: 1.264705408
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Exercises
   :topics: Functions/Exercises
   :from_source: T
   :nocodelens: 
   :pct_on_first: 0.288961039
   :total_students_attempting: 308
   :num_students_correct: 249.0
   :mean_clicks_to_correct: 5.4939759036

   Write a function that removes all occurrences of a given letter from a string.
   ~~~~
   def remove_letter(theLetter, theString):
       # your code here
   
   ====
   
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
           self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
           self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")
   
   
   
   myTests().main()