.. actex:: ac11_14_2
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Exercises
   :topics: Functions/Exercises
   :from_source: T
   :pct_on_first: 0.4811827957
   :total_students_attempting: 372
   :num_students_correct: 351.0
   :mean_clicks_to_correct: 2.1367521368

   Write a function that will return the number of digits in an integer.
   ~~~~
   def numDigits(n):
       # your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
     def testOne(self):
         self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
         self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
         self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
         self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")
   
   
   
   myTests().main()