.. activecode:: writingcode_question10_7
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 10-tuples
   :subchapter: 10-writecode
   :topics: 10-tuples/10-writecode
   :from_source: T
   :practice: T
   :nocodelens:

   Create a tuple called my_data that contains one element, the integer 99.
   ~~~~

   ====
   from unittest.gui import TestCaseGui

   class MyTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(my_data, (99,), "Testing that the tuple has one element, 99.")

   MyTests().main()