.. activecode:: writingcode_question10_1
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 10-tuples
   :subchapter: 10-writecode
   :topics: 10-tuples/10-writecode
   :from_source: T
   :practice: T
   :nocodelens:

   Create a tuple named tup1 that has three elements: 'a', 'b', and 'c'.
   ~~~~

   ====
   from unittest.gui import TestCaseGui

   class MyTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(tup1, ('a', 'b', 'c'), "Checking to make sure the tuple has the correct elements.")

   MyTests().main()