.. activecode:: lhs_sample_activecode
    :author: LHS CS Team
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: GeneralIntro
    :subchapter: Exercises
    :topics: GeneralIntro/Exercises
    :from_source: F
    :language: python
    :pct_on_first: 0.6488549618
    :total_students_attempting: 262
    :num_students_correct: 255.0
    :mean_clicks_to_correct: 1.7215686275

    Write a program that prints out "Hello World!"
    ~~~~
    # your code here
    ====
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def testOne(self):
            self.assertIn('Hello World!', self.getOutput())
    
        def testTwo(self):
            self.assertIn('print', self.getEditorText())
    
    myTests().main()