.. activecode:: ac11_15_1
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: Functions
   :subchapter: ChapterAssessment
   :topics: Functions/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T

   Write a function called ``int_return`` that takes an integer as input and returns the same integer.
   ~~~~
   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(int_return(10), 10, "Testing that function int_return(10) returns 10")

   myTests().main()