.. actex:: atheno_hw1_q2
   :author: Atheno Chen
   :difficulty: 0.0
   :basecourse: fopp
   :chapter: SimplePythonData
   :subchapter: Exercises
   :topics: SimplePythonData/Exercises
   :from_source: F
   :autograde: unittest
   :pct_on_first: 0.5
   :total_students_attempting: 34
   :num_students_correct: 19.0
   :mean_clicks_to_correct: 1.2105263158

   Write a program that will compute the area of a trapezoid. Prompt the user to enter base 1, base 2 and height and save them to variables called ``a``, ``b`` and ``h``. Print a nice message back to the user with the answer.
   ~~~~
   
   
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def test_radius(self):
           self.assertIn("a", self.getEditorText(), 'base 1 variable')
           self.assertIn("b", self.getEditorText(), 'base 2 variable')
           self.assertIn("h", self.getEditorText(), 'height variable')
           self.assertIn(str((a+b)*h/2)[:4], self.getOutput(), 'Checking answer.')
   myTests().main()