.. actex:: ex_2_12
   :author: bmiller
   :difficulty: 1.0714215748
   :basecourse: fopp
   :chapter: SimplePythonData
   :subchapter: Exercises
   :topics: SimplePythonData/Exercises
   :from_source: T
   :pct_on_first: 0.45
   :total_students_attempting: 380
   :num_students_correct: 287.0
   :mean_clicks_to_correct: 2.212543554

   Write a program that will compute MPG for a car. Prompt the user to enter the number of
   miles driven and store it in a variable called `miles` and the number of gallons used stored in a variable `gallons`.
   Print a nice message with the answer.
   ~~~~
   
   ====
   from unittest.gui import TestCaseGui
   import re
   class myTests(TestCaseGui):
       def testOne(self):
          self.assertTrue(re.search(str(miles/gallons)[:5], self.getOutput()), 'Checking answer.')
   myTests().main()