.. actex:: ex_2_14
   :author: bmiller
   :difficulty: 1.1037037037
   :basecourse: fopp
   :chapter: SimplePythonData
   :subchapter: Exercises
   :topics: SimplePythonData/Exercises
   :from_source: T
   :autograde: unittest
   :pct_on_first: 0.3271604938
   :total_students_attempting: 486
   :num_students_correct: 330.0
   :mean_clicks_to_correct: 2.7606060606

   Ask the user for the temperature in Fahrenheit and store it in a variable call ``deg_f``. Calculate the equivalent temperature in degrees Celsius and store it in ``deg_c``. Output a message to the user giving the temperature in Celsius.
   ~~~~
   
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):
           self.assertIn('deg_f', self.getEditorText())
           self.assertIn('deg_c', self.getEditorText())
           ans = (deg_f - 32)*5/9
           self.assertAlmostEqual(ans, deg_c)
           self.assertIn(str(ans)[:min(len(str(ans)), 5)],    self.getOutput(), 'Checking answer.')
   myTests().main()