.. actex:: ex_2_10
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: SimplePythonData
   :subchapter: Exercises
   :topics: SimplePythonData/Exercises
   :from_source: T
   :autograde: unittest
   :pct_on_first: 0.3460526316
   :total_students_attempting: 760
   :num_students_correct: 561.0
   :mean_clicks_to_correct: 2.9590017825

   Write a program that will compute the area of a circle. Prompt the user to enter the radius and save it to avariable called ``radius``. Print a nice message back to the user with the answer.
   ~~~~
   pi = 3.14
   
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def test_radius(self):
           self.assertIn("radius", self.getEditorText(), 'radius variable')
           self.assertIn(str(3.14*radius**2)[:4], self.getOutput(), 'Checking answer.')
   myTests().main()