.. actex:: ex_2_11
   :author: bmiller
   :difficulty: 1.0567525577
   :basecourse: fopp
   :chapter: SimplePythonData
   :subchapter: Exercises
   :topics: SimplePythonData/Exercises
   :from_source: T
   :autograde: unittest
   :pct_on_first: 0.3529411765
   :total_students_attempting: 731
   :num_students_correct: 411.0
   :mean_clicks_to_correct: 1.9635036496

   Challenge: Write a program that will compute the area of a rectangle.  Prompt the user to enter the width and height of the rectangle  and store the values in variables called ``width`` and
   ``height``. Print a nice message with the answer..
   ~~~~
   
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):
           self.assertIn('width', self.getEditorText())
           self.assertIn('height', self.getEditorText())
           self.assertIn(str(width*height), self.getOutput(), 'Checking answer.')
   myTests().main()