.. activecode:: ac8_9_2
   :author: bmiller
   :difficulty: 1.0744671258
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: TheAccumulatorPatternwithLists
   :topics: TransformingSequences/TheAccumulatorPatternwithLists
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.4102362205
   :total_students_attempting: 1270
   :num_students_correct: 965.0
   :mean_clicks_to_correct: 2.2642487047

   2. For each word in the list ``verbs``, add an -ing ending. Save this new list in a new list, ``ing``.
   ~~~~
   verbs = ["kayak", "cry", "walk", "eat", "drink", "fly"]
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testSix(self):
         self.assertEqual(ing, ['kayaking', 'crying', 'walking', 'eating', 'drinking', 'flying'], "Testing that the variable ing has the correct value.")
   
   myTests().main()