.. activecode:: ac8_9_3
   :author: bmiller
   :difficulty: 1.0203353095
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: TheAccumulatorPatternwithLists
   :topics: TransformingSequences/TheAccumulatorPatternwithLists
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.6375212224
   :total_students_attempting: 1178
   :num_students_correct: 924.0
   :mean_clicks_to_correct: 1.3452380952

   Given the list of numbers, ``numbs``, create a new list of those same numbers increased by 5. Save this new list to the variable ``newlist``.
   ~~~~
   numbs = [5, 10, 15, 20, 25]
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testFour(self):
         self.assertEqual(newlist, [10, 15, 20, 25, 30], "Testing that the newlist value contains the correct elements.")
   
   myTests().main()