.. activecode:: lhs_10_3
    :author: LHS CS Team
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: Lists
    :subchapter: Exercises
    :topics: Lists/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 0.1403508772
    :total_students_attempting: 456
    :num_students_correct: 421.0
    :mean_clicks_to_correct: 10.3610451306

    Write a function that returns the sum of all the numbers in a list
    that are even and divisible by 3. To receive full credit, you must
    use a **list comprehension**.
    
    ~~~~
    # My Name:
    
    def sumEvenAndDiv3(lst):
        # write your code here
    
    # any test cases you'd like to have
    
    ====
    import random
    import re
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
        def _sumEvenAndDiv3(self, lst):
            sum = 0
            for e in [num for num in lst if num%2==0 and num%3==0]:
                sum += e
            return sum
        
        def _makeList(self, size, min, max):
            l = []
            for i in range(size):
                l.append(random.randrange(min,max))
            return l
    
        def testOne(self):
            SelfTest = False
            if (SelfTest):
                self.assertEqual(0,1,"Auto-testing if OFF - you must check the program yourself for now")
                return
    
            print('Auto-testing...')
            test = [ [1, 3, 5, 7, 9], [1, 2, 3, 6, 9, 12, 13, 15] ]
            for i in range(len(test)):
                self.assertEqual(sumEvenAndDiv3(test[i]), self._sumEvenAndDiv3(test[i]), "Test "+str(i))
            
            l = self._makeList(15, 0, 30)
            self.assertEqual(sumEvenAndDiv3(l), self._sumEvenAndDiv3(l), "Random List 1")
            editText = self.getEditorText()
            Found = len(re.findall("for .* in .*[^\n].*if", editText)) >= 1
            if Found:
                self.assertEqual(Found, True, "list comprehension used")
            else:
                self.assertEqual(Found, True, "list comprehension not found")
    myTests().main()