.. activecode:: TL_365_h1_q4_1
   :author: Tyler Luchko
   :difficulty: 1.0537284272
   :basecourse: fopp
   :chapter: GeneralIntro
   :subchapter: Exercises
   :topics: GeneralIntro/Exercises
   :from_source: F
   :language: python
   :autograde: unittest
   :pct_on_first: 0.4132231405
   :total_students_attempting: 242
   :num_students_correct: 148.0
   :mean_clicks_to_correct: 1.9121621622

   Initialize two variables to values ``9`` and ``2`` respectively. 
   Then:
    
   1. increment each variable first by value ``2`` (means you 
   update those variables)
    
   2. integer divide the two variables (divide the larger value with 
   the lower one), and assign the result to another variable called 
   ``c``
    
   3. print the variable ``c`` and make sure you have the desired 
   result
   ~~~~
   
   ====	       
   
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):
           self.assertEqual(self.getOutput(), '2\n', 'Checking output.')
           self.assertEqual(c, 2, 'Checking c.')
   myTests().main()