.. actex:: ac14_9_3
   :author: bmiller
   :difficulty: 1.0121269391
   :basecourse: fopp
   :chapter: MoreAboutIteration
   :subchapter: Exercises
   :topics: MoreAboutIteration/Exercises
   :from_source: T
   :pct_on_first: 0.8163265306
   :total_students_attempting: 147
   :num_students_correct: 136.0
   :mean_clicks_to_correct: 1.2058823529

   Using a while loop, create a list called ``nums`` that contains the numbers 0 though 20. (i.e: your while looop should initialize a counter variable on 0. During each iteration, the loop should append the current value of the counter variable to ``nums`` and then increase the counter by 1. The while loop should stop once the counter variable is greater than 20)
   ~~~~
   
   ====
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(nums, [0,1,2,3,4,5,6,7,8,9,10,11,12,13,14,15,16,17,18,19,20],"Testing that nums has been assigned the correct elements")
   
   myTests().main()