.. actex:: ac14_9_2
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: MoreAboutIteration
   :subchapter: Exercises
   :topics: MoreAboutIteration/Exercises
   :from_source: T
   :pct_on_first: 0.7058823529
   :total_students_attempting: 170
   :num_students_correct: 159.0
   :mean_clicks_to_correct: 1.3396226415

   Using a while loop, create a list called ``L`` that contains the numbers 0 to 10. (i.e.: Your while loop should initialize a counter variable to 0. On each iteration, the loop should append the current value of the counter variable to ``L`` and then increase the counter by 1. The while loop should stop once the counter variable is greater than 10.)
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(L, [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10], "Testing that L was created correctly.")
   
   myTests().main()