.. activecode:: ee_ch13_011
   :author: bmiller
   :difficulty: 1.1492190986
   :basecourse: fopp
   :chapter: Classes
   :subchapter: ImprovingourConstructor
   :topics: Classes/ImprovingourConstructor
   :from_source: T
   :tags:Classes/ImprovingourConstructor.rst: 
   :pct_on_first: 0.3167808219
   :total_students_attempting: 584
   :num_students_correct: 435.0
   :mean_clicks_to_correct: 3.5333333333

   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOneA(self):
         self.assertEqual(t.num1, 6, "Testing that t.num1 has correct number assigned.")
      def testOneB(self):
         self.assertEqual(t.num2, 10, "Testing that t.num2 has correct number assigned.")
   
   myTests().main()