.. activecode:: question10_5_3
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 10-tuples
   :subchapter: 10-multipleassignments
   :topics: 10-tuples/10-multipleassignments
   :from_source: T
   :nocodelens:
   :practice: T

   Write code to create a list called 'l' and add the key-value pairs of dictionary d to list l as tuples. Sort list l by the value in descending order.
   ~~~~
   d = {'monkey': 5, 'snake': 3, 'rabbit': 9, 'dragon': 6, 'rooster': 2, 'rat': 10}

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(l, [('rat', 10), ('rabbit', 9), ('dragon', 6), ('monkey', 5), ('snake', 3), ('rooster', 2)], "Making sure the list was sorted correctly.")

   myTests().main()