.. activecode:: MSIT-02-03
    :author: Kevin Austin
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: SimplePythonData
    :subchapter: Exercises
    :topics: SimplePythonData/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 0.4736842105
    :total_students_attempting: 38
    :num_students_correct: 36.0
    :mean_clicks_to_correct: 2.3333333333

    Write a statement to create a float variable named nine that is equal to 9
    ~~~~
    # add your code after this line (requires no more than one statement)
    
    =====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def testOne(self):
           self.assertIsInstance(nine, float, "Testing that nine is a float variable.")           
           self.assertEqual(nine, 9, "Testing that nine is equal to 9.0.")
    
    myTests().main()