.. activecode:: lhs_16_11a
    :author: LHS CS Team
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: ClassesBasics
    :subchapter: Exercises
    :topics: ClassesBasics/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 0.3308270677
    :total_students_attempting: 399
    :num_students_correct: 395.0
    :mean_clicks_to_correct: 3.564556962

    Create a class called ``Cereal`` that has three instance variables: ``name``, ``brand``, and ``fiber``.
    When an instance of ``Cereal`` is printed, the user should see the following:
    "[name] cereal is produced by [brand] and has [fiber] grams of fiber in every serving!"
    To the variable name ``c1``, assign an instance of ``Cereal`` whose name is **Corn Flakes**,
    brand is **Kellogg's**, and fiber is 2. To the variable name ``c2``, assign an instance of
    ``Cereal`` whose name is **Honey Nut Cheerios**, brand is **General Mills**, and fiber is 3.
    Practice printing both! 
    
    ~~~~
    #Define the class here
    
    #Your test code here
    ====
    
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
        def testOne(self):
            SelfTest = False
            if (SelfTest):
                self.assertEqual(0,1,"Auto-testing if OFF - you must check the program yourself for now")
                return
    
        def testOne(self):
            self.assertEqual(c1.__str__(), "Corn Flakes cereal is produced by Kellogg's and has 2 grams of fiber in every serving!", "Testing that c1 prints correctly.")
        def testTwo(self): 
            self.assertEqual(c2.__str__(), "Honey Nut Cheerios cereal is produced by General Mills and has 3 grams of fiber in every serving!", "Testing that c2 prints correctly.")
         
    myTests().main()