.. actex:: ex_7_10
   :author: bmiller
   :difficulty: 3.0
   :basecourse: thinkcspy
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: T
   :practice: T
   :autograde: unittest
   :pct_on_first: 0.5936762624
   :total_students_attempting: 2119
   :num_students_correct: 1999.0
   :mean_clicks_to_correct: 2.2696348174

   Write a function that will return the number of digits in an integer.
   ~~~~
   def numDigits(n):
       # your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
     def testOne(self):
         self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
         self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
         self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
         self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")
   
   
   
   myTests().main()