.. actex:: ex_8_6
   :author: bmiller
   :difficulty: 3.0
   :basecourse: thinkcspy
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: T
   :practice: T
   :autograde: unittest
   :nocodelens: 
   :pct_on_first: 0.4628699051
   :total_students_attempting: 1791
   :num_students_correct: 1670.0
   :mean_clicks_to_correct: 3.3928143713

   def mirror(mystr):
       # your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(mirror("good"),"gooddoog","Tested mirror on input of 'good'")
           self.assertEqual(mirror("Python"),"PythonnohtyP","Tested mirror on input of 'Python'")
           self.assertEqual(mirror(""),"","Tested mirror on input of ''")
           self.assertEqual(mirror("a"),"aa","Tested mirror on input of 'a'")
   
   
   myTests().main()