.. actex:: ex_8_7
   :author: bmiller
   :difficulty: 3.0
   :basecourse: thinkcspy
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: T
   :practice: T
   :autograde: unittest
   :nocodelens: 
   :pct_on_first: 0.320177205
   :total_students_attempting: 2483
   :num_students_correct: 2057.0
   :mean_clicks_to_correct: 5.4919786096

   Write a function that removes all occurrences of a given letter from a string.
   ~~~~
   def remove_letter(theLetter, theString):
       # your code here
   
   ====
   
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
     def testOne(self):
         self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
         self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
         self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")
   
   
   
   myTests().main()