.. actex:: ex_8_10
   :author: bmiller
   :difficulty: 3.0
   :basecourse: thinkcspy
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: T
   :practice: T
   :autograde: unittest
   :nocodelens: 
   :pct_on_first: 0.412837345
   :total_students_attempting: 1371
   :num_students_correct: 1252.0
   :mean_clicks_to_correct: 4.2691693291

   Write a function that removes the first occurrence of a string from another string.
   ~~~~
   def remove(substr,theStr):
       # your code here
   
   ====
   
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
           def testOne(self):
               self.assertEqual(remove("an","banana"),"bana","Tested remove on inputs of 'an' and 'banana'")
               self.assertEqual(remove("cyc","bicycle"),"bile","Tested remove on inputs of 'cyc' and 'bicycle'")
               self.assertEqual(remove("iss","Mississippi"),"Missippi","Tested remove on inputs of 'iss' and 'Mississippi'")
               self.assertEqual(remove("egg","bicycle"),"bicycle","Tested remove on inputs of 'egg' and 'bicycle'")
   
   
   
   myTests().main()