.. actex:: ex_8_11
   :author: bmiller
   :difficulty: 3.0
   :basecourse: thinkcspy
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: T
   :practice: T
   :autograde: unittest
   :pct_on_first: 0.3266850069
   :total_students_attempting: 1454
   :num_students_correct: 1241.0
   :mean_clicks_to_correct: 6.4206285254

   Write a function that removes all occurrences of a string from another string.
   ~~~~
   def remove_all(substr,theStr):
       # your code here
   
   
   
   ====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
     def testOne(self):
         self.assertEqual(remove_all("an","banana"),"ba","Tested remove_all on inputs of 'an' and 'banana'")
         self.assertEqual(remove_all("cyc","bicycle"),"bile","Tested remove_all on inputs of 'cyc' and 'bicycle'")
         self.assertEqual(remove_all("iss","Mississippi"),"Mippi","Tested remove_all on inputs of 'iss' and 'Mississippi'")
         self.assertEqual(remove_all("eggs","bicycle"),"bicycle","Tested remove_all on inputs of 'eggs' and 'bicycle'")
   
   
   
   myTests().main()