.. activecode:: lhs_9_17
    :author: LHS CS Team
    :difficulty: 0.0
    :basecourse: thinkcspy
    :chapter: Strings
    :subchapter: Exercises
    :topics: Strings/Exercises
    :from_source: F
    :language: python
    :autograde: unittest
    :pct_on_first: 0.3882352941
    :total_students_attempting: 85
    :num_students_correct: 74.0
    :mean_clicks_to_correct: 3.3513513514

    Write a function ``printComment(lineOfCode)`` that takes the input string which
    is one line of code and returns a string which is the ``lineOfCode`` from the 
    first ``#`` to the end of the line.
    
    As practice, do not use any string method
    
    ``printComment("def f(x):  # this is a function")`` returns ``# this is a function``
    
    
    ~~~~
    def printComment(lineOfCode):
        # write your code here
    
    # any test cases you'd like to have
    
    ====
    import sys
    import re
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
    
        def _printComment(self, s):
            out = ""
            foundComment = False
            for ch in s:
                if ch == '#':
                    foundComment = True
                if foundComment:
                    out += ch
            return out
    
        # just opening and closing file does not work, so need a print statement to "empty" the file
        def deleteFile(self, fname):
            # empty out the file so nothing shows
            sys.stdout = open(fname, "w")
            print()
            sys.stdout.close()
    
        def testOne(self):
            print("Auto-testing...")  # can be deleted if you hide all prints() to a file
            
            # strips the test code - when testing code with the test code 
            code = self.getEditorText()
            pos = code.find("class myTests")
            if (pos != -1):
                code = code[:pos]
    
            # hide print()
            fname = "test.out"
            # create output file
            sys.stdout = open(fname, "w")
    
            tests = [ 'adf adfa ad fad # asdf ad  ',
                      '# asdfa # asdfad # ',
                      '   asdf   ###    adfadsf' ]
            for t in tests:
                self.assertEqual(printComment(t), self._printComment(t), t)
    
            sys.stdout.close()
            self.deleteFile(fname)
            
            
    myTests().main()