.. actex:: rab_lq3_ex1
    :author: Richard Bernatz
    :difficulty: 1.0
    :basecourse: thinkcspy
    :chapter: Strings
    :subchapter: Exercises
    :topics: Strings/Exercises
    :from_source: F
    :nocodelens: 
    :pct_on_first: 0.1279069767
    :total_students_attempting: 86
    :num_students_correct: 50.0
    :mean_clicks_to_correct: 7.92

    Use this active code window for coding problem 1. Include the required function only!
    
    
    ~~~~
    # your name
    def order_add_cut(s1, s2):
    
    
    ====
    from unittest.gui import TestCaseGui
    
    class myTests(TestCaseGui):
        def testOne(self):
            self.assertTrue(order_add_cut('dog', 'cat') == 'atdo', "Tested 'dog', 'cat'")
            self.assertTrue(order_add_cut('zebra', 'wildabeast') == 'ildabeastzebr', "Tested 'zebra', 'wildabeast'")
            self.assertTrue(order_add_cut('trails', 'trail') == 'railtrail', "Tested 'trails', 'trail'")
            self.assertTrue(order_add_cut('Monday', 'Friday') == 'ridayMonda', "Tested 'Monday', 'Friday'")
           
          
    
    myTests().main()