.. activecode:: ac15_2_4
   :author: bmiller
   :difficulty: 1.0282916443
   :basecourse: fopp
   :chapter: AdvancedFunctions
   :subchapter: KeywordParameters
   :topics: AdvancedFunctions/KeywordParameters
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.7875457875
   :total_students_attempting: 273
   :num_students_correct: 254.0
   :mean_clicks_to_correct: 1.4803149606

   **6.** Currently the function is supposed to take 1 required parameter, and 2 optional parameters, however the code doesn't work. Fix the code so that it passes the test. This should only require changing one line of code.
   ~~~~
   
   def waste(var = "Water", mar, marble = "type"):
       final_string = var + " " + marble + " " + mar
       return final_string
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(waste("Pokemon"), "Water type Pokemon", "Testing that waste returns the correct string on input 'Pokemon'")
   
   myTests().main()