.. activecode:: ac5_6_5
   :author: bmiller
   :difficulty: 1.1497897261
   :basecourse: fopp
   :chapter: Sequences
   :subchapter: TheSliceOperator
   :topics: Sequences/TheSliceOperator
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.1354115928
   :total_students_attempting: 2053
   :num_students_correct: 1569.0
   :mean_clicks_to_correct: 3.5430210325

   Create a new list using the 9th through 12th elements (four items in all) of ``new_lst`` and assign it to the variable ``sub_lst``.
   ~~~~
   new_lst = ["computer", "luxurious", "basket", "crime", 0, 2.49, "institution", "slice", "sun", ["water", "air", "fire", "earth"], "games", 2.7, "code", "java", ["birthday", "celebration", 1817, "party", "cake", 5], "rain", "thunderstorm", "top down"]
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(sub_lst, new_lst[8:12], "Testing that sub_lst has the correct elements assigned.")
   
   myTests().main()