.. activecode:: ac5_9_5
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: Sequences
   :subchapter: SplitandJoin
   :topics: Sequences/SplitandJoin
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.4799789252
   :total_students_attempting: 1898
   :num_students_correct: 1456.0
   :mean_clicks_to_correct: 2.1923076923

   Create a variable ``output`` and assign to it a list whose elements are the words in the string ``str1``.
   ~~~~
   str1 = "OH THE PLACES YOU'LL GO"
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testSix(self):
         self.assertEqual(output, ["OH", "THE", "PLACES", "YOU'LL", "GO"], "Testing that output value is assigned to correct value.")
   
   myTests().main()