.. activecode:: ac12_4_9
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: Tuples
   :subchapter: TupleAssignmentwithunpacking
   :topics: Tuples/TupleAssignmentwithunpacking
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :chatcodes: 
   :pct_on_first: 0.5848563969
   :total_students_attempting: 383
   :num_students_correct: 305.0
   :mean_clicks_to_correct: 1.4

   With only one line of code, assign the variables ``water``, ``fire``, ``electric``, and ``grass`` to the values "Squirtle", "Charmander", "Pikachu", and "Bulbasaur"
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(water, "Squirtle", "Testing that water is assigned to the correct value.")
         self.assertEqual(fire, "Charmander", "Testing that fire is assigned to the correct value.")
         self.assertEqual(electric, "Pikachu", "Testing that electric is assigned to the correct value.")
         self.assertEqual(grass, "Bulbasaur", "Testing that grass is assigned to the correct value.")
   
   myTests().main()