.. actex:: JOC7_14_6
    :author: Joaquin Carbonara
    :difficulty: 0.0
    :basecourse: fopp
    :chapter: Conditionals
    :subchapter: Exercises
    :from_source: F
    :practice: T
    :topics: Conditionals/TheAccumulatorPatternwithConditionals
    :nocodelens:

    Provided is a list of numbers. For each of the numbers in the list, use conditional statements (i.e. "if") to determine whether they are even. If the number is even, add ``True`` to a new list called ``is_even``. If the number is odd, then add ``False``.
    ~~~~
    num_lst = [3, 20, -1, 9, 10]

    ====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):
         def testOne(self):
             self.assertEqual(is_even, [False, True, False, False, True],"Testing that is_even is set correctly.")

    myTests().main()