.. activecode:: ex_6_9
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Selection
   :subchapter: Exercises
   :topics: Selection/Exercises
   :from_source: None

   def isLeap(y):
       # type your code from exercise 8 here


   def daysInMonth(month, year):
       # your code goes here


   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):
       def testOne(self):
           self.assertEqual(daysInMonth(1,2001),31,"January")
           self.assertEqual(daysInMonth(3,2002),31,"March")
           self.assertEqual(daysInMonth(4,2003),30,"April")
           self.assertEqual(daysInMonth(5,2004),31,"May")
           self.assertEqual(daysInMonth(6,2005),30,"June")
           self.assertEqual(daysInMonth(7,2006),31,"July")
           self.assertEqual(daysInMonth(8,2007),31,"August")
           self.assertEqual(daysInMonth(9,2008),30,"September")
           self.assertEqual(daysInMonth(10,2009),31,"October")
           self.assertEqual(daysInMonth(11,2010),30,"November")
           self.assertEqual(daysInMonth(12,2011),31,"December")
           self.assertEqual(daysInMonth(2,2000),29,"February - leap year")
           self.assertEqual(daysInMonth(2,2012),29,"February - leap year")
           self.assertEqual(daysInMonth(2,2013),28,"February - not leap year")
           self.assertEqual(daysInMonth(2,2100),28,"February - not leap year")

   myTests().main()