.. activecode:: ex_6_5
    :author: jenkins
    :difficulty: 3.0
    :basecourse: msumoorhead
    :chapter: Selection
    :subchapter: Exercises
    :topics: Selection/Exercises
    :from_source: None

    def is_even(n):
        # your code here

    ====

    from unittest.gui import TestCaseGui
    import random
    class myTests(TestCaseGui):
         def testOne(self):
             r = random.randrange(0,101,2)
             self.assertEqual(is_even(r),True,"Tested on input of "+str(r))
             r = random.randrange(1,101,2)
             self.assertEqual(is_even(r),False,"Tested on input of "+str(r))
             self.assertEqual(is_even(1),False,"Tested on input of 1")
             self.assertEqual(is_even(0),True,"Tested on input of 0")

    myTests().main()