.. activecode:: ex_6_8
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Selection
   :subchapter: Exercises
   :topics: Selection/Exercises
   :from_source: None

   def isLeap(year):
       # your code here

   ====
   from unittest.gui import TestCaseGui
   import random

   class myTests(TestCaseGui):
       def testOne(self):
           self.assertEqual(isLeap(1800),False,"Tested isLeap on an input of 1800")
           self.assertEqual(isLeap(1900),False,"Tested isLeap on an input of 1900")
           self.assertEqual(isLeap(2000),True,"Tested isLeap on an input of 2000")
           y = 2000
           while y == 2000:
               y = random.randint(490,510) * 4
           self.assertEqual(isLeap(y),True,"Tested isLeap on an input of "+str(y))
           y = random.randint(490,510) * 4 + 1
           self.assertEqual(isLeap(y),False,"Tested isLeap on an input of "+str(y))
           y = random.randint(490,510) * 4 + 2
           self.assertEqual(isLeap(y),False,"Tested isLeap on an input of "+str(y))
           y = random.randint(490,510) * 4 + 3
           self.assertEqual(isLeap(y),False,"Tested isLeap on an input of "+str(y))
           self.assertEqual(isLeap(2100),False,"Tested isLeap on an input of 2100")


   myTests().main()