.. activecode:: ex_6_7
    :author: jenkins
    :difficulty: 3.0
    :basecourse: msumoorhead
    :chapter: Selection
    :subchapter: Exercises
    :topics: Selection/Exercises
    :from_source: None

    def is_even(n):
        # type your code from exercise 5 here

    def is_odd(n):
        # your code here

    ====
    from unittest.gui import TestCaseGui
    import random
    class myTests(TestCaseGui):
         def testOne(self):
             r = random.randrange(1,101,2)
             self.assertEqual(is_odd(r),True,"Tested on input of "+str(r))
             r = random.randrange(0,101,2)
             self.assertEqual(is_odd(r),False,"Tested on input of "+str(r))
             self.assertEqual(is_odd(1),True,"Tested on input of 1")
             self.assertEqual(is_odd(0),False,"Tested on input of 0")

    myTests().main()