.. activecode:: ac9_4_3
   :author: bmiller
   :difficulty: 1.1007620768
   :basecourse: fopp
   :chapter: Files
   :subchapter: AlternativeFileReadingMethods
   :topics: Files/AlternativeFileReadingMethods
   :from_source: T
   :available_files: emotion_words2.txt
   :language: python
   :nocodelens: 
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.4012158055
   :total_students_attempting: 1316
   :num_students_correct: 947.0
   :mean_clicks_to_correct: 2.7106652587

   3. Create a string called ``first_forty`` that is comprised of the first 40 characters of ``emotion_words2.txt``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(first_forty, 'Sad upset blue down melancholy somber bi', "Testing that first_forty was created correctly.")
   myTests().main()