.. activecode:: ac9_4_1
   :author: bmiller
   :difficulty: 1.1893763734
   :basecourse: fopp
   :chapter: Files
   :subchapter: AlternativeFileReadingMethods
   :topics: Files/AlternativeFileReadingMethods
   :from_source: T
   :language: python
   :nocodelens: 
   :autograde: unittest
   :practice: T
   :available_files: school_prompt2.txt
   :pct_on_first: 0.2278566599
   :total_students_attempting: 1479
   :num_students_correct: 1060.0
   :mean_clicks_to_correct: 4.2150943396

   1. Using the file ``school_prompt2.txt``, find the number of characters in the file and assign that value to the variable ``num_char``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(num_char, 537, "Testing that num_char has the correct value.")
   
   myTests().main()