.. activecode:: ac9_4_2
   :author: bmiller
   :difficulty: 1.0662865261
   :basecourse: fopp
   :chapter: Files
   :subchapter: AlternativeFileReadingMethods
   :topics: Files/AlternativeFileReadingMethods
   :from_source: T
   :available_files: travel_plans2.txt
   :language: python
   :nocodelens: 
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.4011670314
   :total_students_attempting: 1371
   :num_students_correct: 1029.0
   :mean_clicks_to_correct: 2.1253644315

   2. Find the number of lines in the file, ``travel_plans2.txt``, and assign it to the variable ``num_lines``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testTwo(self):
         self.assertEqual(num_lines, 11, "Testing that num_lines is assigned to correct value.")
   
   myTests().main()