.. activecode:: ac9_4_1
   :author: bmiller
   :difficulty: 1.1895551991
   :basecourse: fopp
   :chapter: Files
   :subchapter: AlternativeFileReadingMethods
   :topics: Files/AlternativeFileReadingMethods
   :from_source: T
   :language: python
   :nocodelens: 
   :autograde: unittest
   :practice: T
   :available_files: school_prompt2.txt
   :pct_on_first: 0.2273342355
   :total_students_attempting: 1478
   :num_students_correct: 1059.0
   :mean_clicks_to_correct: 4.2181303116

   1. Using the file ``school_prompt2.txt``, find the number of characters in the file and assign that value to the variable ``num_char``.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(num_char, 537, "Testing that num_char has the correct value.")
   
   myTests().main()