.. activecode:: question9_4_2
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 09-dictionaries
   :subchapter: 09-loopinganddictionaries
   :topics: 09-dictionaries/09-loopinganddictionaries
   :from_source: T
   :practice: T
   :nocodelens:

   Write code to add the keys of dictionary wordCount into a list named MoreThan5 if the keys' value is greater than 5.
   ~~~~
   wordCount = {'so': 7, 'the': 10, 'doctor': 2, 'there': 4, 'code': 10, 'a': 27, 'TI-84': 2, 'therefore': 5}

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(len(MoreThan5), 4, "Testing to make sure the length of the list changed.")

   myTests().main()