.. actex:: ac21_6_12
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: AdvancedAccumulation
   :subchapter: Exercises
   :topics: AdvancedAccumulation/Exercises
   :from_source: T
   :pct_on_first: 0.3777777778
   :total_students_attempting: 45
   :num_students_correct: 38.0
   :mean_clicks_to_correct: 2.5789473684

   Write a function called ``longlengths`` that returns the lengths of those strings that have at least 4 characters. Try it using map and filter.
   ~~~~
   
   def longlengths(strings):
       return None
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(longlengths(["Hello", "hi", "bye", "wonderful"]), [5, 9], "Testing whether longlengths has been correctly defined.")
           self.assertIn('map(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
           self.assertIn('filter(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
   
   myTests().main()