.. actex:: ac21_6_9
   :author: bmiller
   :difficulty: 1.040670619
   :basecourse: fopp
   :chapter: AdvancedAccumulation
   :subchapter: Exercises
   :topics: AdvancedAccumulation/Exercises
   :from_source: T
   :pct_on_first: 0.6
   :total_students_attempting: 45
   :num_students_correct: 42.0
   :mean_clicks_to_correct: 1.6904761905

   Define longwords using filter.
   ~~~~
   
   def longwords_Fil(strings):
       """Return a shorter list of strings containing only the strings with more than four characters. Use the filter function."""
       # write your code here
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(longwords_Fil(["Hello", "hi", "bye", "wonderful"]), ["Hello", "wonderful"], "Testing whether longwords_Fil has been correctly defined.")
           self.assertNotIn('map(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
           self.assertIn('filter(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
   
   myTests().main()