.. actex:: ac21_6_2
   :author: bmiller
   :difficulty: 1.1163319946
   :basecourse: fopp
   :chapter: AdvancedAccumulation
   :subchapter: Exercises
   :topics: AdvancedAccumulation/Exercises
   :from_source: T
   :pct_on_first: 0.4489795918
   :total_students_attempting: 49
   :num_students_correct: 40.0
   :mean_clicks_to_correct: 2.975

   Use manual accumulation to define the lengths function below.
   ~~~~
   
   def lengths(strings):
       """lengths takes a list of strings as input and returns a list of numbers that are the lengths
       of strings in the input list. Use manual accumulation!"""
       # fill in this function's definition to make the test pass.
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
       def testOne(self):
           self.assertEqual(lengths(["Hello", "hi", "bye"]), [5, 2, 3], "Testing whether lengths has been correctly defined.")
           self.assertNotIn('map(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
           self.assertNotIn('filter(', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
   
   myTests().main()