.. activecode:: ex_8_8
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: None
   :nocodelens:

   def is_pal(myStr):
       # your code here

   ====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

   def testOne(self):
       self.assertEqual(is_pal("robert"),False,"Tested is_palindrome on input of 'robert'")
       self.assertEqual(is_pal("bob"),True,"Tested is_palindrome on input of 'bob'")
       self.assertEqual(is_pal("racecar"),True,"Tested is_palindrome on input of 'racecar'")
       self.assertEqual(is_pal("starrats"),True,"Tested is_palindrome on input of 'starrats'")
       self.assertEqual(is_pal(""),True,"Tested is_palindrome on input of ''")




   myTests().main()