.. activecode:: ex_8_5
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: None
   :nocodelens:

   def reverse(astring):
       # your code here

   ====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
         self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
         self.assertEqual(reverse(""),"","Tested reverse on input of ''")




   myTests().main()