.. activecode:: ex_8_19
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Strings
   :subchapter: Exercises
   :topics: Strings/Exercises
   :from_source: None

   def remove_dups(astring):
       # your code here


   print(remove_dups("mississippi"))   #should print misp

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def testOne(self):
         self.assertEqual(remove_dups("pooh"),"poh","Tested remove_dups on string 'pooh'")
         self.assertEqual(remove_dups("mississippi"),"misp","Tested remove_dups on string 'mississippi'")
         self.assertEqual(remove_dups("potato"),"pota","Tested remove_dups on string 'potato'")
         self.assertEqual(remove_dups("bookkeeper"),"bokepr","Tested remove_dups on string 'bookkeeper'")
         self.assertEqual(remove_dups("oo"),"o","Tested remove_dups on string 'oo'")

   myTests().main()