.. activecode:: ac8_10_2
   :author: bmiller
   :difficulty: 1.1046714674
   :basecourse: fopp
   :chapter: TransformingSequences
   :subchapter: TheAccumulatorPatternwithStrings
   :topics: TransformingSequences/TheAccumulatorPatternwithStrings
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.3004149378
   :total_students_attempting: 1205
   :num_students_correct: 749.0
   :mean_clicks_to_correct: 2.7770360481

   1. For each character in the string already saved in the variable ``str1``, add each character to a list called ``chars``.
   ~~~~
   str1 = "I love python"
   # HINT: what's the accumulator? That should go here.
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testTwo(self):
         self.assertEqual(chars, ['I', ' ', 'l', 'o', 'v', 'e', ' ', 'p', 'y', 't', 'h', 'o', 'n'], "Testing that chars is assigned to correct values.")
         self.assertIn('append', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
   
   myTests().main()