.. activecode:: ac10_6_2
   :author: bmiller
   :difficulty: 1.0580159214
   :basecourse: fopp
   :chapter: Dictionaries
   :subchapter: AccumulatingResultsFromaDictionary
   :topics: Dictionaries/AccumulatingResultsFromaDictionary
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.3516988062
   :total_students_attempting: 1089
   :num_students_correct: 731.0
   :mean_clicks_to_correct: 1.9849521204

   **1.** The dictionary ``travel`` contains the number of countries within each continent that Jackie has traveled to. Find the total number of countries that Jackie has been to, and save this number to the variable name ``total``. Do not hard code this!
   ~~~~
   travel = {"North America": 2, "Europe": 8, "South America": 3, "Asia": 4, "Africa":1, "Antarctica": 0, "Australia": 1}
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(total, 19, "Testing that total is correct.")
   
   myTests().main()