.. activecode:: ex_7_23
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: MoreAboutIteration
   :subchapter: Exercises
   :topics: MoreAboutIteration/Exercises
   :from_source: None

   def pow2(n):
       # type your code here


   ====
   from unittest.gui import TestCaseGui
   import random
   import math
   class myTests(TestCaseGui):

       def testOne(self):
           p = random.randint(9,150)
           self.assertEqual(pow2(p),2**int(math.ceil(math.log(p,2))),"Tested on "+str(p))
           q = p
           while q == p:
               q = random.randint(9,150)
           self.assertEqual(pow2(q),2**int(math.ceil(math.log(q,2))),"Tested on "+str(q))
           r = p
           while r == p or r == q:
               r = random.randint(9,150)
           self.assertEqual(pow2(r),2**int(math.ceil(math.log(r,2))),"Tested on "+str(r))
           s = 2 ** random.randint(3,10)
           self.assertEqual(pow2(s),s*2,"Tested on "+str(s))

   myTests().main()