.. activecode:: ac17_6_1
   :author: bmiller
   :difficulty: 3.0
   :basecourse: fopp
   :chapter: NestedData
   :subchapter: ChapterAssessment
   :topics: NestedData/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :topics: NestedData/ListswithComplexItems

   The variable ``nested`` contains a nested list. Assign 'snake' to the variable ``output`` using indexing.

   ~~~~

   nested = [['dog', 'cat', 'horse'], ['frog', 'turtle', 'snake', 'gecko'], ['hamster', 'gerbil', 'rat', 'ferret']]

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         self.assertEqual(output, "snake", "Testing that output is assigned to correct value")

   myTests().main()