.. actex:: hollins_table_intro_1_ac
   :author: Stephen Wassell
   :difficulty: 0.0
   :basecourse: fopp
   :chapter: Iteration
   :subchapter: Exercises
   :topics: Iteration/Exercises
   :from_source: F
   :language: python
   :autograde: unittest
   :pct_on_first: 1.0
   :total_students_attempting: 82
   :num_students_correct: 82.0
   :mean_clicks_to_correct: 1.0

   
   A 2-dimensional table of numbers (also called a matrix) can be 
   realized in Python with a list of lists. In the code segment below,
   a table with 3 rows and 4 columns has been defined, with the 
   assignment statement split over three lines for readability.
   Run the code to see what the table looks like when printed.
        
   ~~~~
   table = [[1, 2, 3, 4],
            [5, 6, 7, 8],
            [9, 10, 11, 12]]
   
   print(table)
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):  
           self.assertTrue(True)
   myTests().main()