.. activecode:: question10_1_3
   :author: bmiller
   :difficulty: 3.0
   :basecourse: py4e-int
   :chapter: 10-tuples
   :subchapter: 10-tuplesareimmutable
   :topics: 10-tuples/10-tuplesareimmutable
   :from_source: T
   :practice: T
   :nocodelens:

   Write code that replaces the third and sixth elements of tuple t with their capitalized versions.
   ~~~~
   t = ('a', 'b', 'c', 'd', 'e', 'f', 'g', 'h')

   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
           self.assertEqual(t, ('a', 'b', 'C', 'd', 'e', 'F', 'g', 'h'), "Making sure 'C' and 'F' are capitalized.")

   myTests().main()