.. activecode:: ac15_5_1
   :author: bmiller
   :difficulty: 1.0241920061
   :basecourse: fopp
   :chapter: AdvancedFunctions
   :subchapter: ChapterAssessment
   :topics: AdvancedFunctions/ChapterAssessment
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.702020202
   :total_students_attempting: 198
   :num_students_correct: 168.0
   :mean_clicks_to_correct: 1.4107142857

   Create a function called ``mult`` that has two parameters, the first is required and should be an integer, the second is an optional parameter that can either be a number or a string but whose default is 6. The function should return the first parameter multiplied by the second.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(mult(2), 12, "Testing that mult returns the correct value on input (2)")
         self.assertEqual(mult(5,3), 15, "Testing that mult returns the correct value on input (3,5)")
         self.assertEqual(mult(4,"hello"), "hellohellohellohello", "testing that mult returns the correct value on input (4, 'hello'")
   
   myTests().main()