.. actex:: TL_365_r5_ex2
   :author: Tyler Luchko
   :difficulty: 1.0
   :basecourse: fopp
   :chapter: Sequences
   :subchapter: Exercises
   :topics: Sequences/Exercises
   :from_source: F
   :language: python
   :autograde: unittest
   :pct_on_first: 0.0588235294
   :total_students_attempting: 136
   :num_students_correct: 72.0
   :mean_clicks_to_correct: 6.9583333333

   Given the list in activecode below write the code to do the
   following:
   
   1. print 5th element from the list ('?') using list indexing,
   2. using list indexing, take elements from the list ls to create a
      new variable called ``message`` containing "who run the world?
      girls", and
   3. Capitalize the first word of the message variable so that it
      prints "Who run the world? girls".
   
   ~~~~
   ls = ['who', 'world', 'run', 'lights', '?', 'girls', 'the', 'sea']
   #your code here
   
   ====
   
   from unittest.gui import TestCaseGui
   import re
   class myTests(TestCaseGui):
       def testOne(self): 
           output = self.getOutput().split('\n')
    editor = self.getEditorText().split('\n')
           self.assertEqual(output[0], '?', 'Checking first line')
    self.assertEqual(message, 'Who run the world? girls', 'Checking if assigned variable gives correct value')
    self.assertFalse(re.search(r'(\?|Who run the world? girls)', '\n'.join(editor[1:])), 'Checking for hardcoding')
   myTests().main()