.. activecode:: ac12_2_2
   :author: bmiller
   :difficulty: 1.0062936678
   :basecourse: fopp
   :chapter: Tuples
   :subchapter: TuplePacking
   :topics: Tuples/TuplePacking
   :from_source: T
   :language: python
   :autograde: unittest
   :chatcodes: 
   :practice: T
   :pct_on_first: 0.7816901408
   :total_students_attempting: 426
   :num_students_correct: 365.0
   :mean_clicks_to_correct: 1.1068493151

   **3.** Create a tuple named ``tup1`` that has three elements: 'a', 'b', and 'c'.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(tup1, ('a', 'b', 'c'), "Testing that tup1 was created correctly.")
   
   myTests().main()