.. activecode:: ex_5_14
   :author: jenkins
   :difficulty: 3.0
   :basecourse: msumoorhead
   :chapter: Functions
   :subchapter: Exercises
   :topics: Functions/Exercises
   :from_source: None

   def sumTo(n):
       # your code here

   ====
   from unittest.gui import TestCaseGui
   import random
   class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(sumTo(0),0,"Tested on input 0")
          self.assertEqual(sumTo(1),1,"Tested on input 1")
          m = random.randint(3,30)
          self.assertAlmostEqual(sumTo(m),m*(m+1)/2,0,"Tested input "+str(m))
          n = m
          while n == m:
              n = random.randint(3,30)
          self.assertAlmostEqual(sumTo(n),n*(n+1)/2,0,"Tested input "+str(n))

   myTests().main()