.. activecode:: ac11_4_7
   :author: bmiller
   :difficulty: 1.0157358471
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Returningavaluefromafunction
   :topics: Functions/Returningavaluefromafunction
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.7008474576
   :total_students_attempting: 1180
   :num_students_correct: 962.0
   :mean_clicks_to_correct: 1.2671517672

   **11.** Write a function called ``change`` that takes one number as its input and returns that number, plus 7.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(change(5), 12,"Testing the function change with input 5")
         self.assertEqual(change(-10), -3, "Testing the function change with input -10")
   
   myTests().main()