.. activecode:: ac11_4_4
   :author: bmiller
   :difficulty: 1.0723639735
   :basecourse: fopp
   :chapter: Functions
   :subchapter: Returningavaluefromafunction
   :topics: Functions/Returningavaluefromafunction
   :from_source: T
   :language: python
   :autograde: unittest
   :practice: T
   :pct_on_first: 0.5007824726
   :total_students_attempting: 1278
   :num_students_correct: 1002.0
   :mean_clicks_to_correct: 2.2285429142

   **8.** Write a function named ``same`` that takes a string as input, and simply returns that string.
   ~~~~
   
   =====
   
   from unittest.gui import TestCaseGui
   
   class myTests(TestCaseGui):
   
      def testOne(self):
         self.assertEqual(same('hello'), 'hello', "Testing the same function on input 'hello'.")
   
   myTests().main()